# Basic Workflow

```python
# Always have your imports at the top
import pandas as pd
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import Imputer
from sklearn.ensemble import RandomForestClassifier
from sklearn.base import TransformerMixin
```

# Immutability

Never change your inputs!

Write a function that drops the first column of a dataframe and then returns the result. Make sure that it doesn't change the input dataframe!

<!-- grade_id: cell-6c8218e1615228f0, points: 0, grade: false, solution: true -->
```python
def drop_column(_df, colname):
    # first copy the dataframe just to be safe
    # YOUR CODE HERE
    raise NotImplementedError()
    # now drop the column without using inplace=True
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-d15af97ae329a075, points: 2, grade: true, solution: false -->
```python
df = pd.DataFrame({
    'hello': [1, 2, 3],
    'world': [4, 5, 6]
})

drop_column(df, 'hello')
drop_column(df, 'world')

assert 'hello' in df.columns
assert 'world' in df.columns
```

# Use functions!


```python
# given this dataframe
df = pd.DataFrame({
    'hello': [1, 2, 3],
    'world': [4, 5, 6]
})
```

```python
# this cell
df['hello'] = df['hello'] * 10
df['world'] = df['hello'] * 5
```

```python
# and this cell
df['hello'] = df['hello'] * 100
df['world'] = df['hello'] * 30
```

<!-- grade_id: cell-5d190a8fbb5c1762, points: 0, grade: false, solution: true -->
```python
# create a function called multiply_by that abstracts away
# the need for duplicate code

# remember to respect immutability!
def multiply_by(_df, hello_multiply, world_multiply):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c9bbe4853895e71a, points: 4, grade: true, solution: false -->
```python
df = pd.DataFrame({
    'hello': [1, 2, 3],
    'world': [4, 5, 6]
})

new_df = multiply_by(df, 10, 15)

assert new_df.hello.equals(df.hello * 10)
assert new_df.world.equals(df.world * 15)

new_df = multiply_by(df, 100, 150)

assert new_df.hello.equals(df.hello * 100)
assert new_df.world.equals(df.world * 150)
```

# scikit pipelines

Make a simple pipeline that
1. Drops all columns that start with the string `evil`
1. Fills all nulls with the median

<!-- grade_id: cell-e9663ebd68ad68f8, points: 0, grade: false, solution: true -->
```python
# Create a pipeline step called RemoveEvilColumns
# YOUR CODE HERE
raise NotImplementedError()


# Create an pipeline using make_pipeline
# 1. removes evil columns
# 2. imputes with the mean
# 3. has a random forest classifier as the last step

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-87b055898e579a1a, points: 14, grade: true, solution: false -->
```python
X = pd.DataFrame({
    'evil_1': ['X'] * 100,
    'evil_2': ['Y'] * 100,
    'not_so_evil': list(range(0, 100))
})
y = pd.Series([x % 2 for x in range(0, 100)])

pipeline.fit(X, y)

assert pipeline.steps[0][0] == 'removeevilcolumns'
assert pipeline.steps[1][0] == 'imputer'
assert pipeline.steps[2][0] == 'randomforestclassifier'
```