In this ocasion we are going to use a dataset containing information for Airbnb rooms in Lisbon. The target variable (we want to predict) is the numerical variable `price`. Thus, this is a regression problem.

```python
import pandas as pd
import numpy as np
np.random.seed()
import random
random.seed()
```

```python
airbnb = pd.read_csv("data/airbnb.csv")
```

```python
airbnb.head()
```

Output:
```
room_id  host_id        room_type       neighborhood  reviews  \
0     6499    14455  Entire home/apt              Belém        8   
1    17031    66015  Entire home/apt           Alvalade        0   
2    25659   107347  Entire home/apt  Santa Maria Maior       63   
3    29248   125768  Entire home/apt  Santa Maria Maior      225   
4    29396   126415  Entire home/apt  Santa Maria Maior      132   

   overall_satisfaction  accommodates  bedrooms  price  
0                   5.0             2       1.0   57.0  
1                   0.0             2       1.0   46.0  
2                   5.0             3       1.0   69.0  
3                   4.5             4       1.0   58.0  
4                   5.0             4       1.0   67.0
```

We clean the dataset. First we remove the columns we dont need (the host_id and room_id have no predictive power for us in this scenario.

```python
airbnb = airbnb.drop(["room_id", "host_id"], axis=1)
```

Now we replace null values in numerical columns by the column mean.

```python
numerical_columns = airbnb.select_dtypes([int, float]).columns
airbnb = airbnb.replace({col:airbnb[col].mean() for col in numerical_columns})
```

Now we turn the categorical variables (text) into numerical with `get_dummies`.

```python
airbnb = pd.get_dummies(airbnb)
```

```python
airbnb.head()
```

Output:
```
reviews  overall_satisfaction  accommodates  bedrooms  price  \
0        8                   5.0             2       1.0   57.0   
1        0                   0.0             2       1.0   46.0   
2       63                   5.0             3       1.0   69.0   
3      225                   4.5             4       1.0   58.0   
4      132                   5.0             4       1.0   67.0   

   room_type_Entire home/apt  room_type_Private room  room_type_Shared room  \
0                          1                       0                      0   
1                          1                       0                      0   
2                          1                       0                      0   
3                          1                       0                      0   
4                          1                       0                      0   

   neighborhood_Ajuda  neighborhood_Alcântara            ...             \
0                   0                       0     
```

### 1.Create a gridsearch with a Decision Tree Regressor using this hyperparameters space:

- max_depth: range(2, 10)
- max_features: range(3, 10)

### Set the scoring function as the MAE (use the scikit-learn scorer)

*Note: Name the grid search `"grid_search"` (without the quotes) so it can be validated*

<!-- grade_id: cell-7db2107d016b5313, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE
```

<!-- grade_id: cell-f17ae330eedfcd89, points: 5, grade: true, solution: false -->
```python
from sklearn.model_selection import GridSearchCV

evaluation_space =  grid_search_parameter_space = {
    'max_depth': range(2, 10),
    'max_features': range(3, 10)
}

assert evaluation_space == grid_search.param_grid
assert type(grid_search) == GridSearchCV
assert grid_search.scoring == "neg_mean_absolute_error"
```

### 2. Adjust it to the data, and save the score for the best model to the variable `grid_search_tree_best_score`

<!-- grade_id: cell-26e8979bed7b5896, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE
```

<!-- grade_id: cell-1b5ba1a7272ba05d, points: 5, grade: true, solution: false -->
```python
# we cannot ensure GridSearch's KfoldCV deterministic behavior. I run the code 1000 times to set 
# a reasonable range
assert -45< grid_search_tree_best_score<-35
```

### 3. Create a randomsearch with a KNN regressor and a random search distribution with the following hyperparameter distribution. Set the random_state to 42, and the number of iterations to 100:

- Number of neighbours between 3 and 20
- distante weight `uniform` or `weight`

*Note: Name the random search `"random_search"` so it can be validated*


<!-- grade_id: cell-277007cae66b40cc, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE
```

<!-- grade_id: cell-38229db7eeb31f44, points: 5, grade: true, solution: false -->
```python
from sklearn.neighbors import KNeighborsRegressor
from scipy.stats import randint
evaluation_space = {
                   "n_neighbors": randint(3, 20),
                   "weights": ["uniform", "distance"]
                  }
for key in evaluation_space:
    assert type(random_search_parameter_space_dist[key]) == type(evaluation_space[key])
    
assert random_search.random_state == 42
assert random_search.n_iter == 100
assert type(random_search.estimator == KNeighborsRegressor.__class__)
```

### 4. Adjust the KNN random search to the data, and save the score for the best model to the variable `random_search_knn_best_score`, and the best parameters to the variable `random_search_knn_best_params`.

<!-- grade_id: cell-958ce717623e71c5, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE
```

<!-- grade_id: cell-aed988faddbb1e35, points: 5, grade: true, solution: false -->
```python
import math
assert math.isclose(random_search_knn_best_params, random_search_knn_best_params, abs_tol=0.1)
```

**BONUS EXERCISE (not graded)**. Try to beat section #2 best decision tree estimator by expanding the hyperparameter search. Name the new Random Search `random_search_expanded`.

<!-- grade_id: cell-28a28f2f99459c59, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER HERE
```

<!-- grade_id: cell-eef4c2e74d87cc54, points: 0, grade: true, solution: false -->
```python
baseline_performance = grid_search.best_score_

assert random_search_expanded.best_score_ > grid_search.best_score_
```