# SLU17 - Data Sufficiency and Selection

In this notebook you will cover the following:

- Feature Importance
- Single Factor Analysis
- Learning Curves

```python
import pandas as pd
import numpy as np
import sklearn
from sklearn.datasets import make_classification
from sklearn.model_selection import train_test_split, learning_curve
from sklearn.metrics import roc_auc_score
from matplotlib import pyplot as plt
from base64 import b64encode # just for grading purposes

%matplotlib inline
```

```python
# load up a classification dataset

X = pd.read_csv('data/exercise_X.csv')
y = pd.read_csv('data/exercise_y.csv')['label']
# give X a quick look
X.head()
```

Output:
```
feature_0  feature_1  feature_2  feature_3  feature_4  feature_5  \
0   1.588078   0.078159  -3.352761  -0.456998  -1.537050  -1.540990   
1   0.600816   0.862691  -0.282597   2.212205  -0.167844  -0.530334   
2   1.322777   1.686777  -2.026925   2.669834  -0.977835  -2.636982   
3  -0.617705   0.211531   1.131231   0.441285  -0.237977   0.147819   
4   1.368048   0.847212  -2.562404   1.410269   0.123509  -1.814067   

   feature_6  feature_7 feature_8  id  
0   0.697395  -2.263426   Meerkat   0  
1   0.344190   0.116734   Meerkat   1  
2  -1.325553   0.698066   Meerkat   2  
3   0.901299   1.146880   Meerkat   3  
4  -0.561118  -1.243117   Meerkat   4
```

```python
# looks like a balanced binary target
y.value_counts()
```

Output:
```
0    502
1    498
Name: label, dtype: int64
```

# Find the first obviously useless feature

Can you determine which of the features contains all uniques and therefore cannot have any predictive power?

```python
# Use this cell to determine which of the features serves as a categorical
# feature and contains all uniques

```

<!-- grade_id: cell-a584b7a530c36ed0, points: 0, grade: false, solution: true -->
```python
# set the variable feature_all_unique to the name of the feature
# that contains all uniques
feature_all_unique = None

# YOUR CODE HERE
feature_all_unique = 'id'
```

<!-- grade_id: cell-805f4867b353fe96, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert b64encode(feature_all_unique.encode()) == b'aWQ='
### END TESTS
```

# Find the second obviously useless feature

This one doesn't contain all uniques but based upon some Single Factor Analysis you should be able to determine which feature isn't worth
bothering with.

```python
# use this cell to do some more SFA on other features to determine
# which of them is useless

```

<!-- grade_id: cell-f99bab2575793069, points: 0, grade: false, solution: true -->
```python
# Use this cell to determine the other obviously useless feature
other_useless_feature = None

# YOUR CODE HERE
other_useless_feature = 'feature_8'
```

<!-- grade_id: cell-6964c9cd17f443b8, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert b64encode(other_useless_feature.encode()) == b'ZmVhdHVyZV84'
### END TESTS
```

<!-- grade_id: cell-949c10443e311151, points: 0, grade: false, solution: true -->
```python
# now drop the features that you determined to be useless

# YOUR CODE HERE
X = X.drop(['id', 'feature_8'], axis=1)
```

<!-- grade_id: cell-476d75f82723fb0c, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert b64encode(str(sorted(X.columns)).encode()) == b'WydmZWF0dXJlXzAnLCAnZmVhdHVyZV8xJywgJ2ZlYXR1cmVfMicsICdmZWF0dXJlXzMnLCAnZmVhdHVyZV80JywgJ2ZlYXR1cmVfNScsICdmZWF0dXJlXzYnLCAnZmVhdHVyZV83J10='
### END TESTS
```

# Find the rest of the useless features

Single Factor Analysis isn't likely to do much in helping us to determine
which of the rest of the features are useless. We'll need to use feature_importances in order to find the rest of these bad boys

<!-- grade_id: cell-2f4e8d74d1f01c95, points: 0, grade: false, solution: true -->
```python
# in order to use feature importances you will need to import a tree
# based algorithm. Let's use a decision tree classifier

# import your decision tree classifier here the exact same way that
# it is done in the learning notebooks

# YOUR CODE HERE
from sklearn.tree import DecisionTreeClassifier
```

<!-- grade_id: cell-cbc4c686952738f8, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert b64encode(str(DecisionTreeClassifier).encode()) == b'PGNsYXNzICdza2xlYXJuLnRyZWUudHJlZS5EZWNpc2lvblRyZWVDbGFzc2lmaWVyJz4='
### END TESTS
```

<!-- grade_id: cell-6fbdfd50396e632f, points: 0, grade: false, solution: true -->
```python
# Now let's train the classifier and get the feature importances

# Create your classifier, assign it to the clf variable and then
# train it on X and y
clf = None

# once the classifier is trained, set the feature importances here
# make it a pandas series with the index being the column names
# so that we can visualize the results
feature_importances = None

# set the random_state=1 and max_depth=5 or the tests won't pass!
# YOUR CODE HERE
clf = DecisionTreeClassifier(random_state=1, max_depth=5)
clf.fit(X, y)
feature_importances = pd.Series(clf.feature_importances_)
```

<!-- grade_id: cell-ca84ae67d444979f, points: 4, grade: true, solution: false -->
```python
### BEGIN TESTS
assert isinstance(clf, sklearn.tree.DecisionTreeClassifier), 'The classifier must be a DecisionTreeClassifier'
assert clf.random_state == 1, 'random_state must be 1'
assert clf.max_depth == 5, 'max_depth must be 5'
assert feature_importances.equals(pd.Series(clf.feature_importances_, index=X.columns)), 'feature_importances must be a pandas series with the X columns as the index'
### END TESTS
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-18-232f389039e1> in <module>()
      3 assert clf.random_state == 1, 'random_state must be 1'
      4 assert clf.max_depth == 5, 'max_depth must be 5'
----> 5 assert feature_importances.equals(pd.Series(clf.feature_importances_, index=X.columns)), 'feature_importances must be a pandas series with the X columns as the index'
      6 ### END TESTS

AttributeError: 'numpy.ndarray' object has no attribute 'equals'
```

```python
feature_importances.plot.barh();
```

<!-- grade_id: cell-0dc20755ad6d11ce, points: 0, grade: false, solution: true -->
```python
# now remove 3 additional useless features

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-47bb1f540f749a88, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert b64encode(str(sorted(X.columns)).encode()) == b'WydmZWF0dXJlXzInLCAnZmVhdHVyZV8zJywgJ2ZlYXR1cmVfNCcsICdmZWF0dXJlXzYnLCAnZmVhdHVyZV83J10='
### END TESTS
```

# The learning curve

Okay now that we have gotten rid of all those useless features, let's focus on getting a sense for how much data we need in order to have
reasonable performance.

<!-- grade_id: cell-f22e3c91b478ba32, points: 0, grade: false, solution: true -->
```python
# Now create a dataframe that has a single feature that is the
# cross validation score in order to help us understand
# how increasing amounts of data affect the performance

# HINT: just use the snippet from the Learning Notebook
cross_val_scores = None

# instantiate a classifier that you will inspect the learning rate of
clf = DecisionTreeClassifier(max_depth=5, random_state=1)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1bded5a2e391786, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
# round in order to compensate for implementation details
expected = [round(x, 3) for x in [
    0.7139696510593271,
    0.7659680638722555,
    0.7669600557920303,
    0.7719741240410745,
    0.7729961763220548
]]

assert cross_val_scores is not None, 'Did you forget to set the cross_val_scores variable?'
assert 'Cross-validation Score' in cross_val_scores.columns, 'Did you forget'
actual = [round(x, 3) for x in cross_val_scores['Cross-validation Score']]
assert bool(actual == expected), 'Are your numbers off?'
### END TESTS
```

```python
cross_val_scores.plot.line(
    title='Decision Tree Learning Curve',
    ylim=(0.7, 0.8)
);
```

<!-- grade_id: cell-fc4e8393bc289ff2, points: 0, grade: false, solution: true -->
```python
# Now select the minimum training set size that this particular classifier
# seems to need in order to maximize it's learning from this dataset.

min_train_set_size = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1669e239252c8492, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert min_train_set_size <= int(b64decode(b'MzAw'))
### END TESTS
```