In this notebook you will practice the following:

- Scatterplots
- Line charts
- Bar charts
- Histograms
- Box plots
- Scaling plots

To learn about data visualization, we are going to use a modified version of [The Movies Dataset](https://www.kaggle.com/rounakbanik/the-movies-dataset) which has information about movies

The dataset is located at `data/movies.csv`, and has the following fields

```
    budget: Movie budget (in $).
    genre: Genre the movie belongs to.
    original_language: Language the movie was originally filmed in.
    production_company: Name of the production company.
    production_country: Country where the movie was produced.
    release_year: Year the movie was released.
    revenue: Movie ticket sales (in $).
    runtime: Movie duration (in minutes).
    title: Movie title.
    vote_average: Average rating in MovieLens.
    vote_count: Number of votes in MovieLens.
    release_year: Year the movie was released
```

```python
import pandas as pd
```

```python
movies = pd.read_csv("data/movies.csv")
```

```python
movies.shape
```

```python
movies.head()
```

Import matplotlib, pyplot and the matplotlib inline magic.

<!-- grade_id: cell-96c602a1da8f81e5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f150c73eacb8515f, points: 2, grade: true, solution: false -->
```python
assert plt
```

Change the default chart size to 10 inches width and 10 inches height

<!-- grade_id: rcParams_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rcParams_test, points: 2, grade: true, solution: false -->
```python

```

<hr>

### Note about the grading

Grading plots is difficult, we are using `plotchecker` to grade the plots with nbgrader. 
For `plotchecker` to work with nbgrader, we need to add on each cell, the line

`axis = plt.gca();`

**After the code required to do the plot**.

For example, if we want to plot a scatter plot showing the relationship between revenue and runtime we would do as follows:

```python
# code required to plot
movies[["budget", "revenue"]].plot.scatter(x="budget",y="revenue" )

# last line in the cell required to "capture" the cell and being able to grade it with nbgrader
axis = plt.gca();
```

<hr>

### How does the budget correlate with the revenue?

<!-- grade_id: budget_revenue_scatter_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: budget_revenue_scatter_test, points: 3, grade: true, solution: false -->
```python

```

### How does the average movie duration evolves over time? Set the plot title to "Average Movie Duration by year" 

<!-- grade_id: duration_year_line_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: duration_year_line_test, points: 3, grade: true, solution: false -->
```python

```

### How does the average revenue vary by movie genre? Label the x-axis as "Average Revenue"

<!-- grade_id: revenue_genre_barh_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

axis = plt.gca();
```

<!-- grade_id: revenue_genre_barh_test, points: 3, grade: true, solution: false -->
```python

```

### How is the variable vote_count distributed? (set the x axis limit to  < 2000)

<!-- grade_id: hist_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

axis = plt.gca();
```

<!-- grade_id: hist_test, points: 3, grade: true, solution: false -->
```python

```

Make a plot that displays the revenue broken by movie language and that allows us to check if there are outliers.

<!-- grade_id: boxplot_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

axis = plt.gca();
```

<!-- grade_id: boxplot_test, points: 4, grade: true, solution: false -->
```python

```


# Ungraded Exercise
Load the file misterious_data.csv and use data visualization to answer the following questions:

* How is the distribution of x in general?
* Are there any outlier in any of the fields?
* Which 2 charts better represent the underlying data?. Change their style to `bmh` and add titles to each chart explaining them 


```python

```