# SLU14 - Algorithms: Exercise notebook

In this notebook we will practice using:

- K-nearest neighbours
- Gaussian Naive Bayes
- Decision Tree
- Random Forest

```python
# Import pandas and numpy
import pandas as pd
import numpy as np

# Import the classifiers we will be using
from sklearn.neighbors import KNeighborsClassifier
from sklearn.naive_bayes import GaussianNB
from sklearn.tree import DecisionTreeClassifier
from sklearn.ensemble import RandomForestClassifier

# Import train/test split function
from sklearn.model_selection import train_test_split

# Import ROC AUC scoring function
from sklearn.metrics import roc_auc_score
```

[The Wisconsin breast cancer dataset](https://www.kaggle.com/uciml/breast-cancer-wisconsin-data) is a breast cancer diagnostic dataset: these features are computed from a digitized image of a fine needle aspirate (FNA) of a breast mass.

`diagnosis` is our target: 0 for benign, 1 for malignant.

First we'll read in our dataset and create the train/test split:

```python
dataset = pd.read_csv('data/breast_cancer.csv', index_col='id')

features = dataset.drop('diagnosis', axis=1)
label = dataset['diagnosis']

X_train, X_test, y_train, y_test = train_test_split(features, label, test_size=0.33, random_state=0)
```

# Exercise 1: Train a K-Nearest-Neighbors Classifier

Set the variable for the instance of the model as `knn`.

<!-- grade_id: cell-ba21e3246562e909, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ada4969bd1800b2d, points: 0, grade: true, solution: false -->
```python
y_pred = knn.predict_proba(X_test)
score = roc_auc_score(y_test, y_pred[:,1])
assert score > 0.7
```

# Exercise 2: Train a Naive Bayes Classifier
Set the variable for the instance of the model as `naive_bayes`.

<!-- grade_id: cell-74e9e6407fc91741, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d83c8fa3b303ca85, points: 0, grade: true, solution: false -->
```python
y_pred = naive_bayes.predict_proba(X_test)
score = roc_auc_score(y_test, y_pred[:,1])
assert score > 0.7
```

# Exercise 3: Train a Decision Tree Classifier
Set the variable for the instance of the model as `decision_tree`.

<!-- grade_id: cell-f6896adae562e4ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-44f85a5b4f8d02aa, points: 0, grade: true, solution: false -->
```python
y_pred = decision_tree.predict_proba(X_test)
score = roc_auc_score(y_test, y_pred[:,1])
assert score > 0.7
```

# Exercise 4: Train a Random Forest Classifier
Set the variable for the instance of the model as `random_forest`.

<!-- grade_id: cell-3662d60aeb6602df, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b62b0b100d4cb42e, points: 0, grade: true, solution: false -->
```python
y_pred = random_forest.predict_proba(X_test)
score = roc_auc_score(y_test, y_pred[:,1])
assert score > 0.7
```