In this ocasion we are going to use a dataset containing information for Airbnb rooms in Lisbon. The target variable (we want to predict) is the numerical variable `price`. Thus, this is a regression problem.

```python
import pandas as pd
import numpy as np
np.random.seed()
import random
random.seed()
```

```python
airbnb = pd.read_csv("data/airbnb.csv")
```

```python
airbnb.head()
```

We clean the dataset. First we remove the columns we dont need (the host_id and room_id have no predictive power for us in this scenario.

```python
airbnb = airbnb.drop(["room_id", "host_id"], axis=1)
```

Now we replace null values in numerical columns by the column mean.

```python
numerical_columns = airbnb.select_dtypes([int, float]).columns
airbnb = airbnb.replace({col:airbnb[col].mean() for col in numerical_columns})
```

Now we turn the categorical variables (text) into numerical with `get_dummies`.

```python
airbnb = pd.get_dummies(airbnb)
```

```python
airbnb.head()
```

### 1.Create a gridsearch with a Decision Tree Regressor using this hyperparameters space:

- max_depth: range(2, 10)
- max_features: range(3, 10)

### Set the scoring function as the MAE (use the scikit-learn scorer)

*Note: Name the grid search `"grid_search"` (without the quotes) so it can be validated*

<!-- grade_id: cell-7db2107d016b5313, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f17ae330eedfcd89, points: 5, grade: true, solution: false -->
```python
assert grid_search.param_grid["max_features"][1] == 4
assert grid_search.param_grid["max_depth"][-1] == 9
assert str(grid_search.__class__) == "<class 'sklearn.model_selection._search.GridSearchCV'>"
assert grid_search.scoring == "neg_mean_absolute_error"
```

### 2. Adjust it to the data, and save the score for the best model to the variable `grid_search_tree_best_score`

<!-- grade_id: cell-26e8979bed7b5896, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1b5ba1a7272ba05d, points: 5, grade: true, solution: false -->
```python
assert -45< grid_search_tree_best_score<-35
```

### 3. Create a randomsearch with a KNN regressor and a random search distribution with the following hyperparameter distribution. Set the random_state to 42, and the number of iterations to 10. Set the scoring to MAE too:

- Number of neighbours between 3 and 20
- distante weight `uniform` or `weight`

*Note: Name the random search `"random_search"` so it can be validated*


<!-- grade_id: cell-277007cae66b40cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-38229db7eeb31f44, points: 5, grade: true, solution: false -->
```python
assert "uniform" in random_search.param_distributions["weights"]    
assert str(random_search.param_distributions["n_neighbors"].__class__) =="<class 'scipy.stats._distn_infrastructure.rv_frozen'>"
assert random_search.random_state == 42
assert random_search.n_iter == 10
assert str(random_search.estimator.__class__) == "<class 'sklearn.neighbors.regression.KNeighborsRegressor'>"
assert random_search.scoring == "neg_mean_absolute_error"
```

### 4. Adjust the KNN random search to the data, and save the score for the best model to the variable `random_search_knn_best_score`, and the best parameters to the variable `random_search_knn_best_params`.

<!-- grade_id: cell-958ce717623e71c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aed988faddbb1e35, points: 5, grade: true, solution: false -->
```python
import math
assert math.isclose(random_search_knn_best_score, -45, abs_tol=1.0)
```